Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
import pandas as pd
import os
import os.path

datadir = "publicdata"
```

**Q1** Read `table6.csv` into a dataframe, then look at the data. Is it in Tidy Data form? Explain your answer.

```python
path = os.path.join(datadir, "table6.csv")
df6 = pd.read_csv(path)
df6
```

<!-- grade_id: cell-b710cf5a6681aa7d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q2** Explore the data provided as `table6.csv`, and then draw the functional dependency.

<!-- grade_id: cell-6b669d3180c0fae1, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q3** This question deals with the tuburculosis dataset. At no point should you rearrange the order of the rows, as we are using the default integer index to perform our checking.

1. Read `table6.csv` into a dataframe `df6`.  Do not specify an index column.
2. Combine 'century' and 'yearDigits' into one column, 'year' (whose values are strings), then drop the two old columns. Use `copy()` to avoid modifying the original data frame. Store the result as `df6a`.
3. Starting from `df6a`, split the column 'rate' into two new columns 'cases' (the number before the slash) and 'population' (the number after). After you're done, drop 'rate'. Store the result as `df6b`.

Also note that our checking assumes the order of operations as given above, so that the column order is consistent between your solution and our own.

<!-- grade_id: cell-f40f1f2f851f4a41, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-704bb13a9f05424c, points: 3, grade: true, solution: false -->
```python
# Testing cell

assert df6.shape == (6,4)
assert df6a.shape == (6,3)
assert df6b.shape == (6,4)
assert df6.iloc[2,3] == '37737/172006362'
assert df6a.iloc[3,2] == '2000'
assert df6b.iloc[4,3] == '1272915272'
assert df6b.iloc[0,2] == "745"
```

**Q4** Read `us_rent_income.csv` into a dataframe, then look at the data. Is it in Tidy Data form? Explain your answer. Note: this dataset contains estimated income and rent in each state, as well as the margin of error for each of these quantities. This data came from the US Census.

```python
path = os.path.join(datadir, "us_rent_income.csv")
df = pd.read_csv(path)
df.head()
```

<!-- grade_id: cell-27e0ebcef6854ded, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q5** Explore the data provided as `us_rent_income.csv`, and then draw the functional dependency.

<!-- grade_id: cell-8fa9d821cc812a75, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q6** Read `us_rent_income.csv` into a dataframe (with "GEOID" as the index), then transform as needed to make it tidy. Store the result as `df_rent`.  If the asserts fail, try to understand the expectation.  Look at the asserts to see the expected row and column of the values, because if you did not define a tidy result, those will not correspond to your solution.

<!-- grade_id: cell-bfe72fc9d49121fb, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c44e986dd0e3d6bb, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert(df_rent.shape == (52,4))
assert(df_rent.iloc[0,0] == 24476.0)
assert(df_rent.iloc[0,1] == 747.0)
assert(df_rent.iloc[0,2] == 136.0)
assert(df_rent.iloc[0,3] == 3.0)
assert(df_rent.iloc[20,0] == 37147.0)
assert(df_rent.iloc[31,1] == 809.0) 
```

**Q7** Read `relig_income.csv` into a dataframe, then look at the data. Is it in Tidy Data form? Explain your answer.

```python
path = os.path.join(datadir, "relig_income.csv")
df = pd.read_csv(path)
df.head()

```

<!-- grade_id: cell-e0ea7df53126956a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q8** Explore the data provided as `relig_income.csv`, and then draw the functional dependency. Note: this data came from a survey of individuals about their religion and their income.

<!-- grade_id: cell-a9b25ca589a1eff5, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q9** Read `relig_income.csv` into a dataframe, then transform as needed to make it tidy. Store the result as `df_rel`. Don't change the order of the rows.

<!-- grade_id: cell-f1e0c971bbb3e44a, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
df_rel.head(20)
```

<!-- grade_id: cell-d9745d93e6cf5727, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert(df_rel.shape == (180,3))
assert(df_rel.iloc[0,0] == "Agnostic")
assert(df_rel.iloc[0,1] == "<$10k")
assert(df_rel.iloc[0,2] == 27)
assert(df_rel.iloc[41,0] == "Evangelical Prot")
assert(df_rel.iloc[89,1] == "$40-50k")
assert(df_rel.iloc[104,2] == 14)

```