Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

Please start by running the cell below, to read in `indicators` from a csv file we have provided. Next, consider reviewing the following tutorials:
- http://pandas.pydata.org/pandas-docs/stable/10min.html
- https://pandas.pydata.org/pandas-docs/stable/tutorials.html

```python
import os
import os.path
import pandas as pd

datadir = "publicdata"

path = os.path.join(datadir, "indicators.csv")
indicators = pd.read_csv(path)
```

```python
print(indicators.head())
print(indicators.tail())
```

**Q1** Solve each part with respect to the `indicators` dataframe, rather than starting with `df1` when forming `df2`, etc.

1. Add a column `gdp_per_capita`, defined from `gdp / pop`, and save the resulting dataframe as `df1`.
2. Use the `drop` method to drop the columns `country` and `gdp`, saving the result as `df2`.
3. Drop the rows corresponding to Aruba in 1961 and 1963. Save the resulting dataframe as `df3`.
4. Update the column `cell` to increase it by 10% in all rows, saving the resulting dataframe as `df4`.

<!-- grade_id: cell-3ee03539117981dc, points: 0, grade: false, solution: true -->
```python
# Solution Cell

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-79ff412f2191639f, points: 4, grade: true, solution: false -->
```python
# Testing cell

assert df1.loc[12535,'gdp_per_capita'] > 0.95
assert df1.loc[12535,'gdp_per_capita'] < 0.97
assert df2.shape == (12540, 5)
assert df3.shape == (12538, 7)
assert df3.loc[2,'pop'] == 0.06
assert df4.loc[12539,'cell'] > 14.16
```

**Q2** In the data directory, you will find two csv files, `educationTop.csv` and `educationBottom.csv`, both based on data hosted by `www.census.gov`. Both have the same columns, and each row is a U.S. metropolitan area, with data on population, education, and unemployment. The first `csv` file contains metropolitan areas starting with A-K, and the second starting with L-Z. Read both into `pandas` data frames, using the column `GEO.id2` as an index. Concatenate these two data frames along the row dimension (with the top one on top), and call the result `educationDF`.

<!-- grade_id: cell-bd76db40512af520, points: 0, grade: false, solution: true -->
```python
# Solution Cell

# YOUR CODE HERE
raise NotImplementedError()
print(educationDF.head().iloc[:,:2])
print(educationDF.tail().iloc[:,:2])
```

<!-- grade_id: cell-b7ea13e0a64fa2d5, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert eduTop.shape == (177,60)
assert eduBot.shape == (191, 60)
assert educationDF.shape == (368, 60)
assert list(eduBot.columns) == list(educationDF.columns)
```

**Q3** In the data directory, you will find two csv files, `educationLeft.csv` and `educationRight.csv`, both based on data hosted by `www.census.gov`. Both have the same rows, and each row is a U.S. metropolitan area, with data on population, education, and unemployment. The first has information on individuals without a college degree, and the second has information on individuals with a college degree. Read both into `pandas` data frames with an index based on column 0.  Concatenate these two data frames along the column dimension, and call the result `educationDF2`.

<!-- grade_id: cell-263678c5d80e9a47, points: 0, grade: false, solution: true -->
```python
# Solution Cell

# YOUR CODE HERE
raise NotImplementedError()

educationDF2.shape
```

<!-- grade_id: cell-7cc8d85dbb14034d, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert educationDF2.shape == (367, 58)

```

**Q4** In the data directory, you will find two csv files, `educationLeftJ.csv` and `educationRightJ.csv`, both based on data hosted by `www.census.gov`. In both, rows represent U.S. metropolitan area, with data on population, education, and unemployment. However, they do not have exactly the same set of rows, and the columns are totally different except for the index column `Geography`. 

0. Read both into `pandas` data frames, with names `educationLeftJ` and `educationRightJ`, with index column appropriately set.
1. Make a copy of `educationLeftJ` called `educationLeftOriginal`.
2. Starting with `educationLeftJ`, do a left join to bring in the data from `educationRightJ`, storing your answer as `dfJ`.
3. Starting with `educationLeftOriginal`, do an inner join to bring in the data from `educationRightJ`, storing your answer as `dfJ2`.
4. Now read the original csv files in as `eduLeft` and `eduRight` with no meaningful index. Then, starting from `eduLeft`, do an inner merge along the column `Geography`, storing your answer as `dfJ3`.

<!-- grade_id: cell-626b87634f06841c, points: 0, grade: false, solution: true -->
```python
# Solution Cell

import os
import os.path
import pandas as pd

datadir = "publicdata"

path = os.path.join(datadir, "indicators.csv")
indicators = pd.read_csv(path)

# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: cell-6650ae6ad1902f4f, points: 4, grade: true, solution: false -->
```python
# Testing cell

assert educationLeftJ.shape == (249, 12)
assert educationLeftOriginal.shape == (249, 12)
assert educationRightJ.shape == (215, 12)
assert dfJ.shape == (249, 24)
assert list(dfJ.columns)[0] == 'Estimate; Less than high school graduate: - In labor force:'
assert dfJ2.shape == (152, 24)
assert dfJ3.shape == (152, 25)
assert list(dfJ3.columns)[0] == 'Geography'
```