Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
import re
import os
import os.path
import io
import sys
import string
from contextlib import redirect_stdout
```

**Q** Warmup. Python strings.  For the following variable names and descriptions, write Python assignment statements to obtain the desired description. The first one is done for you.

Var | Description
----|------------
s1  | The string with characters `Hello world`
s2  | The string with characters `Hello world` plus a trailing newline
s3  | The string with characters `Hello world\n` (so the **literal** backslash followed by the character `n`)
s4  | Using double quotes as the delimiter, a string that starts with a double quote, followed by two backslashes in a row, followed by a newline.

<!-- grade_id: cell-c16682bcdac2ecf2, points: 0, grade: false, solution: true -->
```python
s1 = "Hello world"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fe5417438e8a5e94, points: 2, grade: true, solution: false -->
```python
assert s2[-2] == 'd'
assert s2[-1] == '\n'
assert s3[-3] == 'd'
assert s3[-2] == '\\'
assert s3[-1] == 'n'
assert len(s4) == 4
assert s4[0] == '"'
assert s4[1] == s4[2] == '\\'
assert s4[-1] == '\n'
```

**Q** Consider the following target:

    target = "In 2018, Claire had 1208 and Bob had  932"
    
Suppose we can assume a structure where we have a number following the `had` and some amount of whitespace.  Ultimately, we want to "capture" the two numbers, but not the year, and we want a single pattern that matches multiple times.

Start by assigning to variable `pattern` a "raw" string containing the desired regular expression pattern.  Use the regex101.com site to test and debug your pattern.  Make sure you have a capture group in your pattern that will give the text of the digits on a match.

<!-- grade_id: cell-9727fbe7b94b6328, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q** In the `re` module, as reflected in the table on page 112 of your book, there are *functions* `re.search()`, `re.match()` and `re.fullmatch()` that all return a single *Match Object*.  Each of these takes two arguments, where the first is a pattern and the second is a target string.

[re module documentation](https://docs.python.org/3/library/re.html#module-contents)

For the example developed above, write code to assign to `result1` the result of a `search`, to `result2` the result of a `match`, and to `result3` the result of a `fullmatch`.  Print all three result variables *after you predict what will happen*.  (To predict, you will need to actually read the descriptions in the book table.)

If you predicted correctly, congratulations.  If you predicted incorrectly, check your code and check your understanding of the descriptions of these three functions.

<!-- grade_id: cell-28ca46e9e15ff806, points: 2, grade: true, solution: true -->
```python
target = "In 2018, Claire had 1208 and Bob had  932"
# YOUR CODE HERE
raise NotImplementedError()
```

**Q** Mofify your pattern so that `result3`, the result of invoking the `fullmatch` function, yields a non-None result.  In order to be useful, your new pattern should have *two* capture groups for the two numbers we want to get from the target.

<!-- grade_id: cell-0879653f52286996, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q** With success on the last question, the variable `result3` should refer to a *Match Object*.  Like any object, a match object has attributes and methods.  To invoke a method on an object, we append the variable referring to the object with a dot character followed by the method name and then arguments in parenthesis.

To know the methods of a Match Object, we should look up and refer to the documentation:

[Match Objects](https://docs.python.org/3/library/re.html#match-objects)

While I am helping with a link here, for exercises, you should be able to find the documentation for yourself.

Start by practicing with the `group()` method of a Match object.  Assign the following variables:

Var | Description
----|------------
g0  | The value of the string giving the *entire* string matched in `result3`
g1  | The value of the string resulting from the first capture group of `result3`
g2  | The value of the string resulting from the second capture group of `result3`
g3  | The value of the string giving the *entire* string matched in `result1`
g4  | The value of the string resulting from the first (and only) capture group of `result1`

<!-- grade_id: cell-0a4bc785e17df14d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(g0)
print(g1, g2)
print(g3, g4)
```

<!-- grade_id: cell-a7f08908cfed805f, points: 2, grade: true, solution: false -->
```python
assert g0 == target
assert g1 == '1208'
assert g2 == '932'
assert g4 == '1208'
```

**Q** A significant limitation of `search`, `match`, and `fullmatch` is that all of these functions return **at most** one match.  But even in this limited case, and moreso in the general case, we want to define a pattern with the expectation that they will yield **multiple** matches in a given target.  So for the same target above:

    target = "In 2018, Claire had 1208 and Bob had  932"
    
we want a regular expression that would match `had 1208` and `had  932` as different matches, and with a single capture group that yields the number in each case.

If you passed the last questions asserts, the pattern that you used to obtain `result1` should be the correct pattern, if the full match pattern begins with the `had`.

What we need is an *iterator* that, on each iteration, yields the next *Match Object* for the set of matches of a pattern against a target, and that finishes the iteration when we are out of match targets.

At the same `re` documentation page linked above, find and read the documentation for the `finditer()` function.  You should now be able to "bring it all together" and write a `for` loop that uses a call of `finditer()` as the sequence in the `for`.  In each iteration, print out two things: the full match, and the string for the capture group.

<!-- grade_id: cell-48ccefbf8bbbadf6, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```