Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
import requests
import pandas as pd
import json

import util
from util import *
```

## Set Up

**Edit `creds.json` Before Starting this Homework**

### Record Credentials Information

1. In the same folder as this notebook, right-click on `creds.json` and `Open With` and select `Editor`.
2. In the `tmdb` dictionary entry, put your API key, obtained for your application from TMDb, as the value of the `apikey` key.
3. **Save the File**
4. Run the cell below, which calls a function in the `util` module that simply reads in a paricular sub-dictionary within the creds file, and then assigns a global variable `apikey`. The signature for this function is:

    `read_creds(key, folder=".", file="creds.json")`
    
   where the second and third parameter give the folder and name for the credentials file.  Since we are using a `creds.json` in the current directory, we only need to specify the key for the sub-dictionary we wish to read.

```python
tmdbcreds = read_creds("tmdb")
apikey = tmdbcreds['apikey']
```

**Q1** The reading mentioned the TMDB endpoint `/discover/movie`. Many examples of what can be done with this endpoint are provided at the following link. 

    https://www.themoviedb.org/documentation/api/discover

However, as always with APIs, it is unwise to rely solely on examples. Rather, you must read and understand the API documentation. Navigate the TMDB API documentation and find the url that explains how to use the endpoint endpoint `/discover/movie` (the place with the "Try it out" tab). Store this url as a string called `documentation` below.

<!-- grade_id: cell-dad96c8c9a596585, points: 0, grade: false, solution: true -->
```python
documentation = "" # modify this
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-08eb4b5e826e0678, points: 1, grade: true, solution: false -->
```python
# Testing cell

```

**Q2** Write a function `getGenreDict(my_api_key)` that returns a Python dictionary mapping genre names to id numbers, e.g. `{'Action':28, 'Adventure':12,...}`. Your function should only make a single call to `requests.get()`, with the appropriate endpoint path for TMDB, and with the given API key. As always, return `None` if something goes wrong.

<!-- grade_id: cell-f13033eadf64207e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-c6c9372623d6829a, points: 3, grade: true, solution: false -->
```python
# Testing cell

d = getGenreDict(apikey)
assert d['Action'] == 28
assert d['Animation'] == 16
assert len(d) == 19
assert d['Science Fiction'] == 878

bad = getGenreDict(apikey+'z')
assert bad == None

#print(d)
```

**Q3** Write a function `search_person_id(my_api_key, name)` that uses the `/search/person` endpoint to conduct a search for the given `name`, and returns the `id` of the first entry in the `results` list. This should use only one call to `requests.get()`. As always, return `None` if something goes wrong, and return `-1` if there were no results.

<!-- grade_id: cell-d53a29800554dd8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-5cc129c05c292f70, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert search_person_id(apikey,'tom cruise') == 500
assert search_person_id(apikey+'z','Bill Murray') == None
assert search_person_id(apikey,'Bill Murray') == 1532
assert search_person_id(apikey,'chevy Chase') == 54812
assert search_person_id(apikey,'Humphrey Bogart') == 4110
assert search_person_id(apikey,'Lauren bacall') == 7570
assert search_person_id(apikey,'xzyaxoiasdg') == -1

```

**Q4** You are bored in quarantine and want to generate a list of movies featuring your favorite actors/actresses, and in the genres you like. Suppose `actorList` and `genreList` are lists of strings, like "Humphrey Bogart" or "Adventure". You can use your functions `search_person_id` and `getGenreDict` to find the ids associated with these strings.

After studying the API documentation and examples for the `/discover/movie` endpoint, please write a function `makeBigMovieLoD(my_api_key,actorList,genreList)` that returns a LoD with one dictionary per movie, with columns `title` (that's 'original title'), `id`, `popularity`, `language`, and `overview`. Your function should use a single call to `requests.get()` and should include movies with *all* of the actors on `actorList` and with *any* of the genres on `genreList`. If you come across a non-existant actor or genre (e.g., "abventure"), don't include it in the search. Please sort your results by popularity, from most popular to least popular.

<!-- grade_id: cell-9d69ebbce4ab196f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2d28e8a05a0c3958, points: 4, grade: true, solution: false -->
```python
# Testing cell
LoD = makeBigMovieLoD(apikey,['Humphrey Bogart', 'Lauren Bacall'],['Adventure','Drama','Romance'])
assert LoD[0]['title'] == 'To Have and Have Not'
assert len(LoD) == 2
LoD = makeBigMovieLoD(apikey,['Humphrey Bogart', 'Lauren Bacall'],['Abventure','Romance'])
assert len(LoD) == 1
LoD = makeBigMovieLoD(apikey,['Humphrey Bogart'],['Adventure','Drama','Romance'])
assert len(LoD) == 20

```

**Q5** Write a function `makeBigMovieDF(my_api_key,rating,year)` that returns a dataframe consisting of *all* US movies with the given rating (e.g., G, PG-13, R, etc.) in the given year, sorted by revenue, from largest revenue to smallest. Your columns should be `title`, `release_date`, `popularity`, `vote_average`, and `overview`.

Be careful, as your initial invocation of `requests.get()` might only return Page 1 out of some large number of pages. You should be sure to get the information from all the pages. A helper function would be wise.

Hint: Study the documentation for the `/discover/movie` endpoint carefully. What we have called rating they call something else.

<!-- grade_id: cell-72cc49ca1ca272e7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(get_page('75df0c4f93ec9e484534b8ce1db9a67e','G',4,2010))

```

<!-- grade_id: cell-c9a32ca20216aa2a, points: 2, grade: true, solution: false -->
```python
# Testing cell
df = makeBigMovieDF(apikey,'G',1944)
assert df.shape == (132, 5)
assert 'Bambi' in list(df.iloc[0])

```

<!-- grade_id: cell-9cd434c932b179b6, points: 2, grade: true, solution: false -->
```python
# Testing cell 
df = makeBigMovieDF(apikey,'R',1930)
assert "All Quiet on the Western Front" in list(df.iloc[0])
assert df.shape == (60,5)

# Testing cell (warning: may take time)
#df = makeBigMovieDF('75df0c4f93ec9e484534b8ce1db9a67e','G',2010)
#assert 'Avatar' in list(df.iloc[0])
#assert df.shape == (2832, 5)

```