Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# JSON and Networking

This worksheet is based on sections 2.4 and 19.1 of the textbook.

This in-class activity illustrates some of what you can do with JSON, and then guides you through an activity to think about the basics of network architecture. Both of these topics will be essential for our coming work on HTTP architecture and programming.

Please run the cell below.

```python
import json
import os
```

# JSON Exercises

**Q1** (warmup) In the `babyNames()` function below, open the file named `babynames.csv`. Create a dictionary to contain the names as keys and the numbers as values. Read each line of the file, rather than reading the entire file all at once. Make sure you close the file. Populate and return the dictionary you have created.

Note: if time is of the essence, you can skip this problem.

<!-- grade_id: cell-2cc9348043a8ae46, points: 0, grade: false, solution: true -->
```python
def babyNames():
    """Open the comma-separated value file babynames.csv and return a 
    dictionary with names as keys and numbers of applications as values.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bb25e01796fb8a1c, points: 2, grade: true, solution: false -->
```python
assert isinstance(babyNames(), dict)
assert len(babyNames().keys()) == 6
assert '22127' == babyNames()['Jacob']
```

In case you didn't have time for the problem above, the cell below gives you the dictionary the function above wold have created, which you can use in the problems below.

```python
bN = {"Jacob":"22127", "Ethan":"18002", "Michael":"17350", "Jayden":"17179", "William":"17051", "Alexander":"16756"}

```

**Q2** In the `babyNamesDict()` function below, write the dictionary created using `babynames.csv` into a JSON formatted file without using any of the JSON utilities.  Make sure you close the file. Return the name of the file you have written. Ensure the file type is ".json" and remember that the syntax of a JSON file requires curly braces { } at the beginning and end.

Note: in case time is short, please be aware that the folder already contains the json file you are being asked to create here, as `babynames.json`. If you solve this problem and run your code, it will overwrite that file (hopefully with an identical one, if you solve this correctly!)

<!-- grade_id: cell-cc3e9cd3db0d5784, points: 0, grade: false, solution: true -->
```python
def babyNamesDict():
    """Writes the dictionary from babyNames() out to a file. 
    Returns the name of the file.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-993e67ecfc908374, points: 3, grade: true, solution: false -->
```python
output = babyNamesDict()
assert isinstance(output, str)
assert output[-5:] == '.json'
assert os.path.isfile(output)
```

**Q3** In the `babyNamesDictJSON()` function below, write the dictionary created using `babynames.csv` into a JSON formatted file (e.g. `"babynames2.json"` or `"babynames3.json"`), using the JSON `dump` function. Recall that `dump(data,file)` will dump a dictionary `data` into a JSON `file` (but you first have to `open` the file and use `'w'` to tell Python you want to write). Make sure you close the file after `dump`. Return the name of the file you have written.

Note: if time is short, please be aware that the file `babynames2.json` is already in the folder, so you can complete the next problem even if you don't do this one.

<!-- grade_id: cell-0e0ebb6bd11db4d7, points: 0, grade: false, solution: true -->
```python
def babyNamesDictJSON():
    """Writes the dictionary created by babyNames() into a file named 'babynames2.json using json utilities.
       Returns the name of the file.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-afd6fd68cf5a61cc, points: 2, grade: true, solution: false -->
```python
output2 = babyNamesDictJSON()
#assert output != output2
#assert output2 == "babynames2.json"
assert os.path.isfile(output2)
```

**Q4** In the `compareJSON()` function below, write a utility function that compares the texts of two JSON files given their filenames. Assume that the filenames contain the file extension. Please make use of the `load` function that comes with json files.

<!-- grade_id: cell-cd7f8509023e51eb, points: 0, grade: false, solution: true -->
```python
def compareJSON(filename1, filename2):
    """ Compare the text from two JSON files given their 
    filenames. Return true if they are the same.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c40738447af40f00, points: 2, grade: true, solution: false -->
```python
assert compareJSON("babynames2.json", "babynames2.json")
assert compareJSON(output, output2)
```

# Networking

In order to get some practice with networking, we will use a tool that allows you to trace the route of packets through the internet. Using this tool, you can select a website to communicate with, then send a communication (ping) through the network to talk to this website, then trace the path of the response from the website, as it navigates through routers all over the world.

As you may know, many people and organizations do not like for their data to be tracked as it moves around the web, so for many of the websites you put into this tool, the path will be blocked by a firewall (which prevents the signals traceroute needs to trace the route).

Each time you use the tool to answer a question below, please copy and paste the output so you can reference it later.

**Q5** Please use the link below to determine your IP address and other information that is available at your current location. What if you use a different web browser or "private browsing mode"? If you can, please also check the IP address of another device (e.g. your phone), which you can do by googling "what's my IP address."

https://www.slac.stanford.edu/cgi-bin/nph-traceroute.pl

<!-- grade_id: cell-520d303355c70d17, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q6** Using the same link, click "Yes" to tell the website you do want to ping the Stanford network. Please write down three things you learned from the location information in the pathway. Do you think a LAN was involved?

<!-- grade_id: cell-928601083cb80dc3, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q7** Please copy and paste the pathway you got from the previous problem, then repeat the process to obtain another pathway. Are the two pathways exactly the same? Explain

<!-- grade_id: cell-16ab5b390c344f82, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q8** Now, using the same tool, please type in "personal.denison.edu" and report what you learned from the pathway. What are the endpoints of the pathway? 

<!-- grade_id: cell-bd377dbbf1808802, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q9** Now, using the same tool, please type in "www.whitehouse.gov" and report what you learned from the pathway. Do you think the traceroute got all the way to the White House servers? Justify your answer.

<!-- grade_id: cell-0f384bb5ba33fb17, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q10** Now try the same experiments using this link:

https://www.net.princeton.edu/traceroute.html

For example, when going to Denison, what locations does the tracer go through? Can it make it to the White House? What if you give it the Stanford IP address?

<!-- grade_id: cell-02d38ab216b1b641, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q11** Select one of the intermediate IP addresses from any of your tests above (we know that the intermediate IP addresses did not kill traceroute with a firewall). Then put that IP address into this tool and describe what you learned:

https://gsuite.tools/traceroute

For fun, try to make a path that goes through five countries. You may need to use Google to learn IP addresses for some remote places.

<!-- grade_id: cell-3c1003cc8fd89220, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE