Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# JSON

Please review section 2.4, if needed, and then solve the following problem.

```python
import json
import os
import os.path
datadir = "publicdata"
```

**Q1:** Write a function

    writeSquares(N, filepath)

that creates a list of the squares of integers from 1 to `N` and then writes them out as a JSON-encoded text file to the file location given by `filepath`.  If `N` is less than or equal to 0, no file should be created.  This function returns no value.

<!-- grade_id: cell-7bcc1e2c8df55d0e, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e9787275b7d0de79, points: 2, grade: true, solution: false -->
```python
# Testing cell

filepath = "./squares.json"
writeSquares(5, filepath)
result = json.load(open(filepath, 'r'))
assert result == [1, 4, 9, 16, 25]
writeSquares(1, filepath)
result = json.load(open(filepath, 'r'))
assert result == [1]
os.remove(filepath)
writeSquares(0, filepath)
assert not os.path.isfile(filepath)
```

# Programming with HTTP

## Socket String-Based GET

Please review section 2.4, if needed, and then solve the following two problems.

**Q2** In this question, we want to locate a web-server resource and then use the `mysocket` socket and Python string-based strategy for obtaining the data.

Start by using Chrome or other web browser to navigate to:

    www.discovercs.denison.edu
    
Once there, navigate to the resources associated with Chapter 8, and navigate within the web page to the subsection for Project 8.6.

By hovering, or by clicking and looking at the browser's address line, find the resource that corresponds to "Voting Results 1".

Our objective is to write a function that makes a GET request to obtain that resource:

    GetVotes1()

which returns **just the body** 

Note that the function takes no parameters.  So the function must establish the connection, make the request, obtain the full results, close the connection, parse the results to find where the body starts, and then return the string that point to the end.  Make sure your request uses HTTP 1.1, and includes the `Host` and the `Connection` header lines.

<!-- grade_id: cell-1cdb33106f536a30, points: 0, grade: false, solution: true -->
```python
import mysocket as sock

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-64feab94389c2863, points: 4, grade: true, solution: false -->
```python
s = GetVotes1()
assert s[:7] == 'A B C D' 
assert len(s) == 104
```

**Q3** Enhance your implementation of

    GetVotes1()

so that it returns both the body **and** a dictionary that maps from header-name to value for each of the header lines returned in the result.  You are still required to use the `mysocket` method for issuing requests. Your function should still be called `GetVotes1()` and should begin the same way as your solution to the previous problem. You might consider printing the header line-by-line, and comparing what you see to the assert statements below, to get a sense of how to build the dictionary you are asked for.

<!-- grade_id: cell-130b81d731bc8bf3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-141d91f1611654eb, points: 4, grade: true, solution: false -->
```python
body, headers = GetVotes1()
assert body[:7] == 'A B C D' 
assert len(body) == 104
assert len(headers) == 8
assert 'Date' in headers
assert 'Server' in headers
assert headers['Server'] == "Apache"
assert ':' in headers['Date']
```