Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

## Exercises on Files

### References

In addition to Chapter 2 of the class textbook, some other references that you might find helpful:

- Discovering Computer Science, Jessen Havill, Chapter 6, Section 6.2
- Online CS1 textbook: http://interactivepython.org/runestone/static/thinkcspy/Files/toctree.html

### Exercises

Note that you will lose significant credit if, for each function, you do not provide a **docstring** for the function.  If you do not recall how or where docstrings are to be placed for functions, see, for example:  

- Discovering Computer Science, Jessen Havill, Chapter 3, Section 3.4
- https://docs.python-guide.org/writing/documentation/#writing-docstrings

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "publicdata"
```

**Q:** Write a function that takes a single argument of a file path and, by opening, reading and closing the file, obtains the count of characters in the file.  If the path argument does **not** refer to an existing file, the function should return -1.

    fileCharCount(path)


<!-- grade_id: cell-d8e89dc6e7795f85, points: 0, grade: false, solution: true -->
```python
# Solution cell

def fileCharCount(path):
    count = 0
    # YOUR CODE HERE
    raise NotImplementedError()
    return count
print("Tennyson count:", fileCharCount(os.path.join(datadir, "tennyson.txt")))
```

<!-- grade_id: cell-772314a11ed20586, points: 3, grade: true, solution: false -->
```python
# Testing cell

assert fileCharCount(os.path.join(datadir, "tennyson.txt")) == 124
assert fileCharCount(os.path.join(datadir, "hello.txt")) == 57
assert fileCharCount(os.path.join(datadir, "twolines.txt")) == 25
assert fileCharCount(os.path.join(datadir, "foobar.txt")) == -1
```

Suppose, in the data directory given by python variable `datadir`, there is a file named `tennyson.txt`.  The content of the file consists of the five lines as shown here:

    Twilight and evening bell,  
      And after that the dark!     

    And may there be no sadness of farewell,  
      When I embark;  

So there are two lines, an empty line, and then two following lines.  Each of these five lines is terminated by the newline (`'\n'`).  Some lines begin immediately with a non-space character, and some lines begin with one or more spaces.  Although not visible in the above, some of the lines have space characters at the end, before the terminating newline.


**Q:** For the `tennyson.txt` file, we want to read the lines and return a list of strings, with one entry in the list for each line of the file.  Use repeated invocations of the `readline()` method in your solution, and, because this is specific to the `tennyson.txt` file, your function can use the knowledge that there are five lines in the file.

    tennysonLines()


<!-- grade_id: cell-71057810a2d7ce51, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
print(tennysonLines())
```

<!-- grade_id: cell-a2f54407a10b92cb, points: 3, grade: true, solution: false -->
```python
# Testing cell

filelines = tennysonLines()

assert len(filelines) == 5
assert filelines[0] == 'Twilight and evening bell,  \n'
assert filelines[-1] == '  When I embark;  \n'
assert filelines[2] == '\n'
```

**Q:** Write a function

    lastWords(filepath)

that creates and returns a list of strings consisting of the last word in each of the lines of the text file whose location is given by `filepath`.  The items in the list should not include any newline or whitespace as part of any of the words, and there should not be any items in the list for any lines consisting of all whitespace.  A word, as required here, will include the last sequence of non-whitespace in the line.

<!-- grade_id: cell-09d48197093a7c06, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
filepath = os.path.join(datadir, "tennyson.txt")
print(lastWords(filepath))
```

<!-- grade_id: cell-90f390601d4c01ec, points: 3, grade: true, solution: false -->
```python
# Testing cell

assert lastWords(os.path.join(datadir, "tennyson.txt")) == ['bell,', 'dark!', 'farewell,', 'embark;']
```

**Q:** Sometimes, when processing a file, some number of header, comment, or documentation lines at the beginning of a file must be omitted from the processing and skipped over before processing the remaining lines in the file.  Consider the addition of parameter, `skiplines` with a default value of 0 for the previous problem:

    lineLengths(filepath, skiplines=0)

The function should skip the specified number of lines and then, for the remainder of the lines, compute a list of (line number, line length) tuples, one per processed lines.  The line number advances, even for the skipped lines.  Line lengths should exclude any leading or trailing whitespace (spaces, tabs, trailing newline).  So for the `tennyson.txt` file:

    Twilight and evening bell,  
      And after that the dark!     

    And may there be no sadness of farewell,  
      When I embark;  

the result of `lineLengths("tennyson.txt", skiplines=2)` is the list of tuples: `[(3, 0), (4, 40), (5, 14)]` since lines 1 and 2 are skipped, and line 3 is an empty line, consisting of just a newline.

<!-- grade_id: cell-9662239e79b8a501, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

filepath = os.path.join(datadir, "hello.txt")
print("hello.txt:", lineLengths(filepath))

filepath = os.path.join(datadir, "tennyson.txt")
print("tennyson.txt:", lineLengths(filepath, skiplines=2))
```

<!-- grade_id: cell-3396d1b555aa5e73, points: 5, grade: true, solution: false -->
```python
# Testing cell

filepath = os.path.join(datadir, "hello.txt")
assert lineLengths(filepath) == [(1, 13), (2, 42)]
assert lineLengths(filepath, skiplines=1) == [(2, 42)]
assert lineLengths(filepath, skiplines=2) == []

filepath = os.path.join(datadir, "tennyson.txt")
assert lineLengths(filepath, skiplines=2) == [(3, 0), (4, 40), (5, 14)]
```