Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
import requests
import pandas as pd
import json

import util
from util import *
from IPython.display import Image
```

# 05_HW4

This is the last homework. You will upload to Notebowl the ipynb file, at least two image files, and a CSV file.

## Set Up

**Edit `creds.json` Before Starting this Homework**

### Record Credentials Information

1. In the same folder as this notebook, right-click on `creds.json` and `Open With` and select `Editor`.
2. In the `spotify` dictionary entry, put your Client ID, Client Secret, and also your user id (found here: https://www.spotify.com/us/account/overview/).
3. **Save the File**
4. Run the cell below, which calls a function in the `util` module that simply reads in a particular sub-dictionary within the creds file, and then assigns a global variable `client_id`. The signature for this function is:

    `read_creds(key, folder=".", file="creds.json")`
    
   where the second and third parameter give the folder and name for the credentials file.  Since we are using a `creds.json` in the current directory, we only need to specify the key for the sub-dictionary we wish to read.
   
5. Use the provided `spotify_token.ipynb` to generate tokens as needed to complete this homework (i.e., whenever your token expires, generate a new one and use `spotify_token.ipynb` to update `creds.json`).

<!-- grade_id: cell-fe28bdd52b96d5ed, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q1** Study the documentation for the Spotify endpoint `v1/search`. Follow the procedure from class to obtain a Spotify authorization. Then write a function `search_goo()` that uses the authorization to carry out a search for the string "Goo Goo Dolls" and where the type is 'artists'. Return the resulting Python data structure. When you invoke `requests.get()` it will need both query parameters and headers. Hint: Spotify provides a nice "try it yourself" feature.

<!-- grade_id: cell-8a2897e677b39260, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

#answer = search_goo()

```

<!-- grade_id: cell-26fd6b2371ee5d88, points: 0, grade: true, solution: false -->
```python
# Testing cell

```

<!-- grade_id: cell-bddb139fc8dadd92, points: 4, grade: true, solution: false -->
```python
# Testing cell
assert len(answer) == 1
assert(len(answer['artists'])==7)
assert(answer['artists']['items'][0]['name']=='The Goo Goo Dolls')
```

**Q2** Now we are going to do a POST request to Spotify, to create a new playlist. Please study the documentation at https://developer.spotify.com/console/post-playlists/, specifically which parameters (in both header and body) are required, and what scope is required for an OAuth token to create a public playlist.

1. Log in to your spotify and take a screenshot showing your existing playlists. Modify the cell below to show your screenshot.

<!-- grade_id: cell-f2c876c579348297, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

2. Modify the notebook `spotify_token.ipynb` to generate a token with the appropriate scope. This should update your `creds.json`.
3. Write code below (in a global cell, not a function) that carries out a POST request to create a playlist with the name 'CS 181 05_HW4 Playlist'. You can experiment with the "Try It" button at the API documentation, but in order to pass the test cases you will need to write code below.

<!-- grade_id: cell-8259437a765b3e5d, points: 0, grade: false, solution: true -->
```python
# Solution cell

user_id = "" # change this
token = "" # change this, using spotifycreds


# YOUR CODE HERE
raise NotImplementedError()

results = response.json()
json_head(results,50)
```

<!-- grade_id: cell-461dcf8d61e2f9ed, points: 4, grade: true, solution: false -->
```python
assert results['name'] == 'CS 181 05_HW4 Playlist'
assert results['owner']['id'] == user_id
assert token in headerD['Authorization']
assert 'description' in results
```

4. Take a screenshot from the user side, showing that the client (your code above) successfully created that playlist. Replace the image below with your screenshot.

<!-- grade_id: cell-4fd1b7deed446903, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q3** In this problem you will apply what you've learned in Section 5 of the course. You will select a Web API, obtain any credentials needed, and obtain data from the API. The selected provider must require Authentication and Delegated Authorization (OAuth 2) to enable access to data not available to an anonymous user for at least some subset of the data. You can find a list of APIs here: https://any-api.com/, and you can search for APIs here: http://apis.io/. Please select an API different from those we have already studied, and, by reading through the documentation, please select one that involves OAuth. Then, work through the steps required to get the credentials you need for make a request for data using the `requests` module. Select data you are interested in and read it into a `pandas` dataframe. Then use the `to_csv()` function to write your `pandas` dataframe to a CSV file that you will upload along with your homework. You can read about this function here: https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.to_csv.html

1. In the next cell (a markdown cell), please describe precisely the process you had to do in order to extract data from your chosen API. Please use the text in the book (discussion of TMDB), and in `spotify_token.ipynb` as a template. Make it so that we could replicate your process if we needed to. Include links, and update the `creds.json` document with credentials as you get them. Please also link to the terms of use document, or upload it along with the rest of your homework.

<!-- grade_id: cell-d1fae9a619092e42, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

2. In the next cell, please put code that handles the cells/steps needed to get an access token.  For the most part, these steps only need be executed one time, to go from first registration of a client with a provider to the point where an access token is received (nominally Step 8 of the OAuth Dance).  If tokens need to be refreshed, then the cell(s) needed to refresh a token could be here as well. You can add more cells as needed.

<!-- grade_id: cell-510040fcb09d7f80, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

3. The next cell is a code cell. Please put code there that reads in credentials from `creds.json` (updated by the cell above), then formulates a correct request for data from your chosen API, then reads that data into pandas, and finishes with the `to_csv` command. Your code should start from the premise of a valid token/apikey and do the steps required for acquiring and organizing the data from the provider. The cell should detect and handle errors from the server, including, but not limited to, handling an expired token. 

<!-- grade_id: cell-2956a65c4e71b1ab, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```