Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

## Exercises on Python Native Data Structures

### References

In addition to Chapter 3 of the class textbook, some other references that you might find helpful:

- Discovering Computer Science, Jessen Havill, Sections 8.3 and 9.1
- Online CS1 textbook: https://runestone.academy/runestone/books/published/thinkcspy/Lists/toctree.html and https://runestone.academy/runestone/books/published/thinkcspy/Dictionaries/toctree.html

### Exercises

Note that you will lose significant credit if, for each function, you do not provide a **docstring** for the function.  If you do not recall how or where docstrings are to be placed for functions, see, for example:  

- Discovering Computer Science, Jessen Havill, Chapter 3, Section 3.4
- https://docs.python-guide.org/writing/documentation/#writing-docstrings

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "publicdata"
```

**Q1** Write a function `elementwiseMult(vec1,vec2)` that uses the `zip` function to compute and return the element-wise multiplication of two lists, `vec1` and `vec2`, of the same length. For example, if `vec1 = [2,4,5]` and `vec2 = [4,3,1]`, you would return `[8,12,5]`.

<!-- grade_id: cell-c8d98116207ef509, points: 0, grade: false, solution: true -->
```python
# Solution cell

def elementwiseMult(vec1,vec2):
    # YOUR CODE HERE
    raise NotImplementedError()
print(elementwiseMult([2,4,5],[4,3,1]))
```

<!-- grade_id: cell-c716e4e1e525691d, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert elementwiseMult([2,4,1,5],[4,3,0,1]) == [8,12,0,5]
assert elementwiseMult([-2,5],[4,3]) == [-8,15]

```

**Q2** Write a function

`dropFives(data)`

that creates, accumulates, and returns a new list that has all the elements from the list, `data`, except for the items with integer value 5. The elements in your new list should be in the same order as they were in `data`.

<!-- grade_id: cell-de2912cd0b6a8a98, points: 0, grade: false, solution: true -->
```python
# Solution cell

def dropFives(data):
    # YOUR CODE HERE
    raise NotImplementedError()
print(dropFives([2,4,5,4,3,1]))
```

<!-- grade_id: cell-8db9727f3209420f, points: 1, grade: true, solution: false -->
```python
# Testing cell

assert dropFives([2,4,5,4,3,1]) == [2, 4, 4, 3, 1]
assert dropFives([5,5,5,5]) == []
assert dropFives([5,55,-5,5]) == [55,-5]


```

**Q3** Write a function

`filmography(L)`

whose parameter `L` is a list of tuples of the form (actor, movie), that creates and returns a dictionary where the keys are actor names and the values are lists of movies that actor has been in. For example, given `L = [(‘DiCaprio’,‘Revenant’),(‘Cumberbatch’,‘Hobbit’),(‘DiCaprio’,‘Gatsby’),(‘DiCaprio’,‘Django’)]`, then you should return `{‘DiCaprio’:[‘Revenant’,‘Gatsby’,‘Django’],‘Cumberbatch’:[‘Hobbit’]}`

<!-- grade_id: cell-66fb4389e963a59d, points: 0, grade: false, solution: true -->
```python
# Solution cell

def filmography(L):
    # YOUR CODE HERE
    raise NotImplementedError()
print(filmography([('DiCaprio','Revenant'),('Cumberbatch','Hobbit'),('DiCaprio','Gatsby'),('DiCaprio','Django')]))
```

<!-- grade_id: cell-c4a63419ac36da74, points: 2, grade: true, solution: false -->
```python
# Testing cell

assert filmography([('De Niro','Irishman'),('Cumberbatch','Brexit'),('Hamill','Star Wars'),('DiCaprio','Django')]) == {'De Niro':['Irishman'],'Cumberbatch':['Brexit'],'Hamill':['Star Wars'],'DiCaprio':['Django']}
assert filmography([('De Niro','Irishman'),('Cumberbatch','Brexit'),('DiCaprio','Gatsby'),('DiCaprio','Django')]) == {'De Niro':['Irishman'],'Cumberbatch':['Brexit'],'DiCaprio':['Gatsby','Django']}
assert filmography([]) == {}
assert filmography([('Cumberbatch','Hobbit')]) == {'Cumberbatch':['Hobbit']}

```

**Q4** Write a function

`whichIsBigger(f1,f2,n)`

which computes `f1(n)` and `f2(n)`, and returns whichever function (`f1` or `f2`) results in the larger value when an integer `n` is passed as a parameter. In case of a tie, return `f1`. Test with `f1(x)` defined as x^3 and `f2(x)` defined as 5x^2+4 and using n=2. Test again with n=100.

<!-- grade_id: cell-8b12181a84b0e621, points: 0, grade: false, solution: true -->
```python
# Solution cell

def whichIsBigger(f1,f2,n):
    # YOUR CODE HERE
    raise NotImplementedError()

def f1(x):
    return x**3

def f2(x):
    return 5*x*x + 4

print(whichIsBigger(f1,f2,2))
print(whichIsBigger(f1,f2,100))
```

<!-- grade_id: cell-7c76005eb023b881, points: 1, grade: true, solution: false -->
```python
# Testing cell

def f1(x):
    return x**3 - 7

def f2(x):
    return 5*x*x + 4*x + 2

assert whichIsBigger(f1,f2,0) == f2
assert whichIsBigger(f1,f2,-10) == f2
assert whichIsBigger(f1,f2,2) == f2
assert whichIsBigger(f1,f2,50) == f1

```

**Q5** Write a lambda expression and assign it to the variable `evensquared` such that it squares the value of its single argument if the argument is even, and cubes the argument if it is odd.

Afterward, for example, expression `evensquared(5)` should yield `125`, while `evensquared(4)` should yield `16`.

Hint: We need an expression to handle an `if` type condition.

<!-- grade_id: cell-36f6dfa09775489f, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

print(evensquared(5))
print(evensquared(4))

```

<!-- grade_id: cell-d76bb4b7c8a840cc, points: 1, grade: true, solution: false -->
```python
assert evensquared(5) == 125
assert evensquared(0) == 0
assert evensquared(-1) == -1
assert evensquared(2) == 4
assert evensquared(-2) == 4
```

**Q6** Write a function
```python
word_lengths(s)
```
that, for each whitespace separated word in string `s`, determines the length of the word and returns a list of such lengths, **omitting occurences of the string: `the`**.  For example, `word_lengths("Denison CS is the best")` returns the list `[7, 2, 2, 4]`, with lengths for 'Denison', `'CS'`, `'is'`, and `'best'`, and omitting the occurence of `'the'`.  Your solution MUST use a list comprehension or else you won't receive full credit.

<!-- grade_id: cell-090557ab8470cb32, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(word_lengths("Denison CS is the best"))
```

<!-- grade_id: cell-56f5c219fdcfe89e, points: 2, grade: true, solution: false -->
```python
assert(word_lengths("Denison CS is the best") == [7,2,2,4])
assert(word_lengths("") == [])
assert(word_lengths("a") == [1])
```

```python

```