Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Programming with HTTP

Please start by running the cell below this one.

```python
import requests
```

**Q1** Create a table that describes a `response` object. List the six most common attributes with a brief description of what information a client gets from each (hint: you might consider checking out the inclass from Friday). Then check the online documentation and make a table describing three methods that come with a `response` object, what each does, and when you might use it. You can use the following markdown code for making tables (but copy and paste it into the cell below):

|attribute|description|
|--------|--------------------|
|first-attribute|what it's good for|

|method|description|use case|
|--------|--------------------|----------|
|first-method|what it's good for|when to use it|

<!-- grade_id: cell-3eb005f5ccf91f40, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

**Q2** Since we often need to begin by combining a protocol, a location, and a resource into a string representing a URL, write a function:

`
buildURL(protocol, location, resource)
`

that returns a string URL based on the three component parts of `protocol`, `location`, and `resource`.  Your function should be flexible, so that if a user specifies a protocol as 'http://' or as 'http', you still obtain a valid URL. Hint: slicing can help with this last bit.

<!-- grade_id: cell-41fc6b8d6eb4bf90, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

buildURL('http://', 'httpbin.org', '/get')
```

<!-- grade_id: cell-31c5bc54e442c99a, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert buildURL('http://', 'httpbin.org', '/get') == 'http://httpbin.org/get'
assert buildURL('http', 'httpbin.org', '/get') == 'http://httpbin.org/get'
assert buildURL('https', 'api.kivaws.org', '/v1/loans/newest') == 'https://api.kivaws.org/v1/loans/newest'
```

**Q3** We often want to be able, given a location and a resource, to construct and issue an HTTP request to the URL from the location and resource using the GET HTTP method, and to obtain a **string** consisting of the body of the returned data resource.

Write a function

`
simpleHTTPGet(location, resource)
`

that uses the function `buildURL` from above (with `http` protocol), along with function(s) and attributes from the `requests` module to achieve this functionality.  If, for whatever reason, we do not successfully obtain a valid string data result, the function should detect this and return `None`. This way the caller, if coded properly, can know whether or not the call was successful. Hint: you can find out whether or not `requests.get()` was successful by looking at the value of `status_code`.

<!-- grade_id: cell-598ea6f1cb1a46c2, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
simpleHTTPGet('personal.denison.edu', '/')
```

<!-- grade_id: cell-69d67104c1fb5cd1, points: 3, grade: true, solution: false -->
```python
# Testing cell

s1 = simpleHTTPGet('personal.denison.edu', '/')
assert len(s1) == 464
assert type(s1) is str
s2 = simpleHTTPGet('httpbin.org', '/get')
assert 300 < len(s2) < 350
s3 = simpleHTTPGet('personal.denison.edu', '/~whiteda/DenisonWebsiteInfo.pdf')
assert s3 == None
```

**Q4** You have probably had the experience before of trying to open a webpage, and having a redirect page pop up, telling you that the page has moved and asking if you want to be redirected. The same thing can happen when we write code to make requests. Write a function:

`
getRedirectURL(location, resource)
`

that begins like your function `simpleHTTPGet` but does not allow redirects when invoking `get`. Your function will return a url. If the `get` results in a success status code (please refer to the table "Common HTTP Status Codes" in Chapter 20), you return the original url (obtained from `buildURL`, with `http` protocol). If you detect that `get` tried to redirect, search within the headers to find the `"Location"` it tried to redirect to, and return that URL instead. If you get any other status code, return `None`.

<!-- grade_id: cell-7a364e38a3b112cf, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

print(getRedirectURL('personal.denison.edu','/~kretchmar'))
print(getRedirectURL('personal.denison.edu','/~kretchmar/'))
print(getRedirectURL('personal.denison.edu','/~whiteda/DenisonWebsiteInfo.pdf'))
```

<!-- grade_id: cell-3849e54c14f31757, points: 4, grade: true, solution: false -->
```python
# Testing cell

assert getRedirectURL('personal.denison.edu','/~whiteda') == 'http://personal.denison.edu/~whiteda/'
assert getRedirectURL('personal.denison.edu','/~whiteda/') == 'http://personal.denison.edu/~whiteda/'
assert getRedirectURL('personal.denison.edu','/~whiteda/DenisonWebsiteInfo.pdf') == None

```

**Q5** We saw in class how to add custom headers along with a get request. This is important in a number of contexts, e.g., when you request a large file, it might make sense to ask it to be compressed for transit. This example motivates our assert statements below. Given parallel lists `headerNameList` and `headerValueList`, you can build a dictionary that maps from header names to their associated values (given by the parallel structure). Write a function

`
getCustomHeader(location, resource,headerNameList,headerValueList)
`

that builds a custom header dictionary and then passes it to the `get` method. Your function should call `buildURL` (with `https`) to build the url to pass to `get`. Your function should return the response from the `get` invocation.

<!-- grade_id: cell-8d319aa00876d442, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

r = getCustomHeader('httpbin.org','/get/',['Transfer-Encoding'],['compress'])
request = r.request
print(request.headers)
print()

r = getCustomHeader('httpbin.org','/get/',['Transfer-Encoding','Accept'],['compress','text/html'])
request = r.request
print(request.headers)
print()

r = getCustomHeader('httpbin.org','/get/',[],[])
request = r.request
print(request.headers)

```

<!-- grade_id: cell-82a227c94a4c2151, points: 3, grade: true, solution: false -->
```python
# Testing cell

r = getCustomHeader('httpbin.org','/get/',['Transfer-Encoding'],['compress'])
request = r.request
assert request.headers['Transfer-Encoding'] == 'compress'
assert request.headers['Connection'] == 'keep-alive'
r = getCustomHeader('httpbin.org','/get/',['Transfer-Encoding','Accept'],['compress','text/html'])
request = r.request
assert request.headers['Accept'] == 'text/html'
assert request.headers['Transfer-Encoding'] == 'compress'

```

**Q6** Please write a function

`
postData(location, resource,dataToPost)
`

that uses your `buildURL` function to build a URL (using `https`), then uses the `requests` module to post `dataToPost` to that URL (note: `dataToPost` will be the *body* of the message you send). Please return the `response` returned by the method of `requests` that you invoke. Note: we test your function with `https://httpbin.org/post`, which is set up to allow you to post there.

<!-- grade_id: cell-3e6febf281674186, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

response = postData('httpbin.org','/post',"Wow, what a cool string!")
print(response.status_code)
print(response.request.body)
```

<!-- grade_id: cell-8eb777c810cbcff6, points: 2, grade: true, solution: false -->
```python
# Testing cell

response = postData('httpbin.org','/post',"CS181 is the best")
r = response.request
assert r.method == 'POST'
assert r.body == 'CS181 is the best'
assert r.url == 'https://httpbin.org/post'
```

```python

```