Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
import re

import os
import os.path
import io
import sys
import string
from contextlib import redirect_stdout

datadir = "publicdata"
```

**Q1:** In the directory referenced by `datadir` is a file named `rime.txt`, whose contents are a small prefix of the **Rime of the Ancyent Marinere**, which we read into a string variable, `data`. Write a function

    findAncients(s)

that, using regular expressions and the `re` module, finds **all** instances of "ancient" in the passed string `s`, spelled either with an `'i'` or a `'y'` (i.e. either `'ancient'` or `'ancyent'`), and insensitive to case.  The function should return a list of tuples with the start and end index (inclusive) of the match, or an empty list if no instances were found.  

<!-- grade_id: cell-148b872c82b7fbc1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

path = os.path.join(datadir, "rime.txt")
if not os.path.isfile(path):
    print("Problem finding source file")
    
with open(path, 'r') as f:
    data = f.read()

result = findAncients(data)
if len(result) == 0:
    print("Ancient not found")
else:
    for t in result:
        print(t, data[t[0]:t[1]+1])
```

<!-- grade_id: cell-63d24cafeb3da570, points: 4, grade: true, solution: false -->
```python
text1 = "Does this text match this pattern?"
assert len(findAncients(text1)) == 0

path = os.path.join(datadir, "rime.txt")
if not os.path.isfile(path):
    print("Problem finding source file")
    
with open(path, 'r') as f:
    data = f.read()

assert len(findAncients(data)) == 3
assert findAncients(data)[0][0] == 16
assert findAncients(data)[0][1] == 23
assert findAncients(data)[2][0] == 384
assert findAncients(data)[2][1] == 391
```

**Match and No Match**

Sometimes, regular expression patterns are expected to match certain patterns and, equally important, to **not** match other patterns.  In the next couple exercises, the goal is to find a single pattern such that, in a **list** of targets, the designated targets within the list match and the designated targets fail to match.  In this setup, each possible target is a string without any line termination.

**Q2:** Write a function

    findMatches(pattern, targetlist)

that performs a regular expression search of `pattern` on each of the targets in targetlist.  The function should build a list that, for each target that both successfully finds a match, and the match is the entire string, includes the matched string/target.  For any target that does not match, or where less than the full target is matched, nothing is added to the resultant list.

<!-- grade_id: cell-da4ceb50085b9293, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6e722a0f35d04727, points: 5, grade: true, solution: false -->
```python
targets = ["abc", "abcd", "abcdef"]
assert findMatches(r'abc', targets) == ['abc']
assert findMatches(r'abcd?', targets) == ['abc', 'abcd']
assert findMatches(r'', targets) == []
```

**Q3:** Suppose we want to distinguish between the two columns below:

    Match     | No Match
    --------------------
    pit       | pt
    spot      | Pot
    spate     | peat
    slap two  | part
    respite   | top it

Write and assign to Python variable `pattern` a regular expression that completely matches all the targets in the first column and does **not** match any of the targets in the second column.  Note that, in the asserts, your `findMatches()` function will be invoked to check your solution.  If these asserts fail, the problem could be with your pattern, or could also be a problem with your `findMatches()` function.

<!-- grade_id: cell-9282f4dd87ae28b6, points: 0, grade: false, solution: true -->
```python
pattern = r''
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a1d29edccb29df9d, points: 2, grade: true, solution: false -->
```python
matchlist = ["pit", "spot", "spate", "slap two", "respite"]
nomatchlist = ["pt", "Pot", "peat", "part", "top it"]
matches = findMatches(pattern, matchlist)
assert matches == matchlist
matches = findMatches(pattern, nomatchlist)
assert matches == []
```