Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

Please start by running the cell below, to import `pandas`. Next, consider reviewing the following tutorials:
- http://pandas.pydata.org/pandas-docs/stable/10min.html
- https://pandas.pydata.org/pandas-docs/stable/tutorials.html

```python
import pandas as pd
```

**Q1** The code below defines a List of Lists (plus a header column), but there is some missing data. The data is based on data from [**R for Data Science**](http://r4ds.had.co.nz/) by Garrett Grolemund and Hadley Wickham. 

0. Use `pandas` to read this List of Lists into a dataframe called `df`, with column labels given by `colnames`. 
Then: 
1. Extract the `shape` of the data into a variable called `sh`.
2. Extract the *list* of `cases`, into a variable called `L` (hint: be very careful about type).
3. Define a list of booleans `K` where `K[i]` is `False` exactly when `population[i]` represents missing data. Do this using a `pandas` built-in function, rather than by using the `math` module. The [documentation](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.isna.html) will help.
4. Create a new data frame `df2`, via slicing and `iloc`, corresponding to the rows that are about `Brazil` and the columns for `year` and `cases`. Your data frame should have two rows and two columns.
5. Use a projection or built-in function to define a new dataframe `df3` consisting of only the rows of `df` without missing data. Again, refer to the [documentation](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.isna.html).
6. Use a projection to define a new dataframe `df4` consisting of only columns without missing data. Use `loc`.
7. Create a new data frame `df5` whose row names are of the form `X_year` where `X` is the first letter of the country name, e.g., the first row name should be `A_1999` and the last `C_2001`.

```python
data = [['Afghanistan',  '1999',    745,  19987071, 0],
        ['Afghanistan',  '2000',   2666,  20595360, 1],
        ['Afghanistan',  '2001',   None,  31527618, 0],
        [     'Brazil',  '1999',  37737, 172006362, 2],
        [     'Brazil',  '2000',  80488, 174504898, 0],
        [      'China',  '1999', 212258, 972915272, 1],
        [      'China',  '2000', 213766, 980428583, 2],
        [      'China',  '2001', 215626,      None, 1]]

colnames = ['country', 'year', 'cases', 'population', 'category']
```

<!-- grade_id: cell-56671dc2ed0cbdfb, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-142b562be65aee80, points: 8, grade: true, solution: false -->
```python
# Testing cell

assert df.at[1,'year'] == '2000'
assert df.at[7,'cases'] == 215626.0
assert df.at[2,'category'] == 0
assert df.at[3,'country'] == 'Brazil'
assert list(df.columns) == ['country', 'year', 'cases', 'population', 'category']
assert sh == (8,5)
assert len(L) == 8
assert 2666.0 in L
assert type(L) == list
assert type(K) == list
assert K[2] == False
assert K[6] == True
assert df2.at[3,'year'] == '1999'
assert df2.shape == (2,2)
assert df2.at[4,'cases'] == 80488.0
assert df3.shape == (6,5)
assert df3.at[1,'year'] == '2000'
assert df3.at[6,'cases'] == 213766.0
assert df3.at[3,'country'] == 'Brazil'
assert df4.shape == (8,3)
assert df4.at[1,'year'] == '2000'
assert df4.at[6,'category'] == 2
assert df4.at[3,'country'] == 'Brazil'
assert df5.shape == (8,5)
assert list(df5.index) == ['A_1999', 'A_2000', 'A_2001', 'B_1999', 'B_2000', 'C_1999', 'C_2000', 'C_2001']

```

**Q2** Using the same data as in the previous problem, and working with `df3` to avoid missing data:

1. Create a new Series `cases_per_capita` obtained by dividing `cases` by `population` in `df3`, then add this to `df3` as a new column named `'cpc'`.
2. Use `agg` to compute the `min` and `max` for `cpc`, in that order, and save the result as `a2`.
3. Use `agg` to compute the minimum population and the maximum number of cases, storing the result as a *data frame* named `a3`. Hint: this means you need square brackets when using `agg`.
4. Use `groupby` to group `df3` by `country`, then use `agg` to compute the `mean` for `'population'` and `'cases'` for each group, storing the result as a *data frame* named `a4`.

<!-- grade_id: cell-013fd91bbb300e31, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-a09db6c81faa799e, points: 4, grade: true, solution: false -->
```python
# Testing cell

assert list(cases_per_capita)[0] > 0.00003
assert list(cases_per_capita)[0] < 0.00004
assert df3.at[1,'year'] == '2000'
assert df3.at[6,'cases'] == 213766.0
assert df3.at[4,'category'] == 0
assert df3.at[3,'cpc'] > 0.0002
assert df3.at[3,'cpc'] < 0.00023
assert df3.shape == (6,6)
assert a2.shape == (2,1)
assert a2.at['max','cpc'] < 0.00047
assert a2.at['max','cpc'] > 0.00045
assert a3.shape == (2,2)
assert a3.at['max','cases'] == 213766.0
assert a4.shape == (3,2)
assert (a4.loc['Brazil'])['population']['mean'] == 173255630.0
assert (a4.loc['Afghanistan'])['cases']['mean'] == 1705.5

```