Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Notebook for Practicing Python Cell Type Exercises

On most Python-centric notebooks, there will often be a read-only cell at the top of the notebook that performs a basic set of `import` statements that might be needed along with some global variable definitions for use in the notebook.

This is not to say that, when solving a problem, you would never need to perform your own `import`.  

You should leave this first setup cell as is and execute it before moving on to the cells where you will provide your own solutions to the questions/exercises.  Be sure, when you do so, that you replace **both** the comment line that you should provide a solution, and the Python statement that raises an exception.  Sometimes, in a solution cell, there is code before and/or after the place where you provide your solution. This code should be kept as is, and sometimes gives an appropriate set up and then invocation of a requested function; or the code you provide may be *within* the body of a loop or function, or placed in a context that allows us to capture the result of `print` function invocations.

If you need/want to for your solution, you can add in additional code cells.  But make sure that you also fill in the solution cell, and that you do not change any autograding assert cells.  (You could change them, but the submission and checking process will always change them back.)  Also be careful to ensure that each cell that you expect to be executed is, in fact, a code cell, and not a markdown cell.

You should use the "checking cell", typically with assert statements, as a way of getting feedback on the correctness of your solution.  If the assert statements fail, start by seeing *exactly which assert* failed, and discern the intent and parameters of the test.  This should allow you to start to debug your solution.  Also, if the asserts fail and you think they should pass, **read the problem again**.  Often the issue is not fully satisfying the problem as stated.

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "publicdata"
```

**Q** Write a function

    double(arg)

that doubles the value of its argument, `arg` and returns the result.

<!-- grade_id: cell-314f181e4a1f666c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(double(2.5))
print(double("Tom"))
```

<!-- grade_id: cell-6d70f8c71feec7bb, points: 1, grade: true, solution: false -->
```python
assert double(2.5) == 5.0
assert double("Hey, Boomer") == "Hey, BoomerHey, Boomer"
```

**Q** (Attributed to Havill in Intro textbook) The speeding ticket fine in a nearby town is \\$50 plus \\$5 for each mph over the posted speed limit.  In addition, there is an extra penalty of \\$200 for all speeds above 90 mph.  Write a function

    fine(speedLimit, clockedSpeed)
    
that returns the fined amount (or 0 if `clockedSpeed` $\le$ `speedLimit`).

In this case, the solution cell does not have an example invocation of your code, so you may want to insert a code cell in between the solution cell and the assertion/checking cell for testing your function with various inputs.

<!-- grade_id: cell-2fbeaef44f6597cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d23606ad5eff5fbf, points: 2, grade: true, solution: false -->
```python
assert fine(35, 40) == 75
assert fine(35, 35) == 0
assert fine(35, 36) == 55
assert fine(55, 89.875) == 224.375
assert fine(55, 90) == 225
assert fine(55, 91) == 430
```

**Q** Write code to print a grid of 5 rows with 4 `'*'` characters in each row using nested `for` loops.  This is a case where the solution cell is "capturing" what your code prints, so your solution must replace the comment and the raise exception **at the same level of indentation**.  This allows the capture of the print and checking in the following cell to see if your solution is correct.

<!-- grade_id: cell-e0233023054ab686, points: 0, grade: false, solution: true -->
```python
printed = io.StringIO()
with redirect_stdout(printed):
    # YOUR CODE HERE
    raise NotImplementedError()
print(printed.getvalue())
```

<!-- grade_id: cell-ab718c1fabd4aa2a, points: 2, grade: true, solution: false -->
```python
assert ( printed.getvalue() == '****\n****\n****\n****\n****\n' or
         printed.getvalue() == '****\n****\n****\n****\n****' )
```