Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
# Run this cell first

import pandas as pd
datadir = "publicdata/"
indicators0 = pd.read_csv(datadir+"indicators.csv")
indicators = pd.read_csv(datadir+"indicators.csv", index_col = 'code')
topnames0 = pd.read_csv(datadir+"topnames.csv")
topnames = pd.read_csv(datadir+"topnames.csv", index_col=['year','sex'])

```

**Q1** It was a design decision to make slicing include its *end* row when `index` is used rather than position. Explain on a human level why this is a good design decision.


<!-- grade_id: cell-2ec48337dbf3f8a2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q2** In the generic syntax for a slice, it says *start:end[:stride]*, but the example in the text says `indicators['RUS','CHN',-1]`. Explain.



<!-- grade_id: cell-c3039faa6761125d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q3** How does Python decide which countries should be included in a slice like `indicators['CHN':'RUS']`? Is it always alphabetical? What if someone wanted a slice by GDP?

<!-- grade_id: cell-cc89c90c7b446d75, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q4** Would it ever make sense to introduce a *row* of booleans instead of a *column* of booleans?

<!-- grade_id: cell-4633571697d55431, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q5** How do you think Python extracts the relevant rows when you call `nlargest`? If you were designing `pandas` would you have it sort every time? Justify your answer.

<!-- grade_id: cell-37550a3cc960cfcd, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q6** Write two lines of code that use `loc` to define: 
1. A DataFrame `df1` with the GDPs and populations for China and the USA, from `indicators`.
2. A DataFrame `df2` with all rows of `topnames0` corresponding to male births after 2000.

<!-- grade_id: cell-4ca07052a4097ed4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(len(df1))
print(len(df2))

```

<!-- grade_id: cell-3c5e27cbe84c6dea, points: 1, grade: true, solution: false -->
```python
assert len(df1) == 2
assert len(df2) == 18
assert df2['name'].nunique() == 3
assert df1.shape == (2,2)
```

**Q7** Write three lines of code that iterate over the rows of `topnames` and uses a list accumulator `L` to append counts of male births in years like 1950, 1960, etc. Use an if-statement.

<!-- grade_id: cell-6bf8c0cbc7f4fc29, points: 0, grade: false, solution: true -->
```python
L = []
# YOUR CODE HERE
raise NotImplementedError()
print(len(L))
```

<!-- grade_id: cell-c63815b11691a498, points: 2, grade: true, solution: false -->
```python
assert len(L) == 12
assert L[0] == 9829
assert L[-2] == 34477
```

**Q8** Write two lines of code that iterate over the columns of `indicators` and appends the `max` of every column to an accumulator `K`, using the built-in `max` function. Note that, if the column is all strings, this will take the alphabetically last string.

<!-- grade_id: cell-9a44bee14555e78f, points: 0, grade: false, solution: true -->
```python
K = []
# YOUR CODE HERE
raise NotImplementedError()
print(len(K))
print(K[1])
```

<!-- grade_id: cell-412eb48086337514, points: 2, grade: true, solution: false -->
```python
assert len(K) == 5
assert K[0] == 'Vietnam'
assert K[-2] == 78.54
```

**Q9** Give an example of an overloaded aggregation function/method, and an example of how to call it by either "passing the Series as the argument to the function" or by "invoking a method (with no argument) on a Series object."


<!-- grade_id: cell-80bd77c4893de588, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q10** At a high level, how do you think the method `value_count()` is implemented? Where have you seen this type of problem before?

<!-- grade_id: cell-58b04084f0533f81, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q11** In the example that calls `max`, `median`, and `nunique`, what would happen if you called `nunique` on `topnames0` instead of `topnames0['name']`?

<!-- grade_id: cell-82d10769a97471fd, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q12** Give an example of a dictionary mapping from column names to lists of functions (rather than function names). Use more than just built-in functions. Hint: you don't need parentheses because you are not calling the functions.

<!-- grade_id: cell-c2414e4cc812197a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q13** The command `indicators.max()` applies the function `max` to every column. Write one line of code to perform the `max` aggregation for all the columns except the `country` column, and save the results in `q13`. Use the `agg` method.

<!-- grade_id: cell-f4faf2f9e1fe7dbe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print(len(q13))
```

<!-- grade_id: cell-094faca476024d26, points: 1, grade: true, solution: false -->
```python
assert q13['pop'] == 1386.4
assert q13['cell'] == 1469.88
assert len(q13) == 4
```

**Q14** Use `agg` to return a dataset with the max, min, mean, and median counts (in that order) in `topnames0` for males born in the 1960s. Save your answer in a variable called `q14`


<!-- grade_id: cell-05672656f5628019, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
q14.shape

```

<!-- grade_id: cell-072a1a71ed36afd7, points: 1, grade: true, solution: false -->
```python
assert q14.shape == (4,1)
assert q14['count'][1] == 79990.0
```

**Q15** Write a sentence to explain why `len(groupby_name)` is 18, and why `len(groupby_year)` is 139, in the `topnames` example.

<!-- grade_id: cell-8dd070d297d89cbf, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Q16** Write a line of code that adds a new column `popSize` to `indicators` which takes value 'high' if `pop > 1000`, 'low' if `pop < 100`, and 'medium' otherwise. 


<!-- grade_id: cell-82edbc196735cade, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
indicators.head()
```

<!-- grade_id: cell-6619a20cd039698d, points: 1, grade: true, solution: false -->
```python
assert len(indicators['popSize']) == 5
assert indicators['popSize'][-1] == 'low'
```

**Q17** Building on the question above, use `groupby` to partition `indicators` by this new column `popSize`, storing the result as `groupby_pop`.


<!-- grade_id: cell-4c9225bf393f3305, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
groupby_pop.head()

```

<!-- grade_id: cell-4e7f934a0bd4dc4d, points: 1, grade: true, solution: false -->
```python
assert len(groupby_pop) == 3
assert len(groupby_pop['popSize']) == 3

```

**Q19** Building on the question above, use `agg` to report the average population in each of the three groups.

<!-- grade_id: cell-34a86c27dfc75f2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
q19
```

<!-- grade_id: cell-8dd2d1c7ed7b286f, points: 1, grade: true, solution: false -->
```python
assert len(q19) == 3
assert q19['pop']['mean'][1] == 94.6
```

**Q20** Write one line of code to sort `q19` by `mean`, storing the result in `indicatorsSorted`, and following the model where `namegroup` was sorted. Warning: pay attention to your `inplace` settings.

<!-- grade_id: cell-fd45be7b5c7d66c2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
indicatorsSorted
```

<!-- grade_id: cell-915589165bf6a966, points: 1, grade: true, solution: false -->
```python
assert indicatorsSorted['pop']['mean'][1] == 234.825
```