Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`/`raise NotImplementedError` or "YOUR ANSWER HERE", as well as your name and collaborators below:

# 04_HW2: XPath

> **In the questions that follow, we are looking for XPath declarative solutions to the problems, not procedural solutions.  You will only get 1/2 credit for procedural solutions.**

Please begin by importing whatever modules you need, reading in and parsing the relevant datasets, and familiarizing yourself with them.

**Q1:** Using the provided `bookstore.xml` file, create a Python list called "books" containing the **titles** of all books. Your list `books` should be a list of strings.

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
# Solution cell
books = []
# YOUR CODE HERE
raise NotImplementedError()
print(books)
type(books[0])
```

<!-- grade_id: q3a1, points: 2, grade: true, solution: false -->
```python
assert len(books) > 0 and type(books[0]) is etree._ElementUnicodeResult
assert 'Lover Birds' in books and 'Splish Splash' in books
assert len(books)==12
```

**Q2:** Create a list of books ids named `less` that cost less than `$6`. Note that `id` is an attribute.

<!-- grade_id: q3b, points: 0, grade: false, solution: true -->
```python
# Solution cell
less = []
# YOUR CODE HERE
raise NotImplementedError()
less
```

<!-- grade_id: q3b1, points: 2, grade: true, solution: false -->
```python
assert len(less) > 0 and type(less[0]) is etree._ElementUnicodeResult
assert 'bk104' in less
assert 'bk101' not in less
assert len(less)==7
```

**Q3:** Create a list of book titles called "eva" where Eva Corets was the author. Your list `eva` should be a list of strings.

<!-- grade_id: q3c, points: 0, grade: false, solution: true -->
```python
# Solution cell
eva = []
# YOUR CODE HERE
raise NotImplementedError()
eva
```

<!-- grade_id: q3c1, points: 2, grade: true, solution: false -->
```python
assert len(eva) > 0 and type(eva[0]) is etree._ElementUnicodeResult
assert len(eva)==3
assert 'Maeve Ascendant' in eva
assert 'Paradox Lost' not in eva
```

**Q4:** Find the average book price for all books that are not fantasy in this file, assigning to variable `avgprice`. **Hints** First, use XPath to get a list of the price strings (text) based on a single XPath query.  Then use a list comprehension to build a list of `float` values converting the strings to real-valued numbers.  Finally, perform the average based on the values and length of the list.

<!-- grade_id: q3d, points: 0, grade: false, solution: true -->
```python
# Solution cell
avgprice = 0
# YOUR CODE HERE
raise NotImplementedError()
avgprice
```

<!-- grade_id: cell-431b0c49edcca411, points: 2, grade: true, solution: false -->
```python
assert(avgprice > 23.82)
assert(avgprice < 24)
```

**Q5:** Create a list called `lessFantasy` containing the titles of the books where the price is under `$40` and not in the fantasy genre.

<!-- grade_id: q3e, points: 0, grade: false, solution: true -->
```python
# Solution cell
lessFantasy = []
# YOUR CODE HERE
raise NotImplementedError()
lessFantasy
```

<!-- grade_id: q3e1, points: 2, grade: true, solution: false -->
```python
assert len(lessFantasy)==6
assert 'Paradox Lost' in lessFantasy
assert 'Maeve Ascendant' not in lessFantasy
```

**Q6:** Using `countries.xml`, generate a list of all the countries in the `countries.xml` file, assigning to a variable `countries`; then assign the number of countries to the variable `countrycount`. When you read in and parse the file, please name the root element `croot`.

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
# Solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a99f77ad878ec7b1, points: 2, grade: true, solution: false -->
```python
assert(countrycount == 231)
assert('Uruguay' in countries)
assert type(croot) is etree._Element
```

**Q7:** Write a function `findPop(root,country)` that finds the population of a given `country` in the dataset `countries.xml`. Use an XPath expression and a format string. Return your answer as an integer.

<!-- grade_id: q4c, points: 0, grade: false, solution: true -->
```python
# Solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q4c1, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert findPop(croot,'Cuba') == 10951334
assert findPop(croot,'Uruguay') == 3238952
```

**Q8:** Create a list named "german50" with the **names** of all countries where over 50% of the population speaks German.  This one is more challenging, and may take some additional investigation to do it with just xpath.  Hint1: In a conditional predicate (`[ ]`), we can use `text()` to refer to the text of the current element.  Hint2: Once we get the right "set" of traversals, we can continue to build our path to get the desired set of information.

<!-- grade_id: q4e, points: 0, grade: false, solution: true -->
```python
# Solution cell
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q4e1, points: 2, grade: true, solution: false -->
```python
assert len(german50)==3
assert type(german50[0]) is etree._ElementUnicodeResult
```

**Q9:** Study the `countries` data carefully. Then use the `position()` function to create a node set consisting of, for countries in positions 5-55 inclusive, the population of the second city listed, if there are at least two cities listed. For example, nothing is in the node set for Aruba (no cities listed) or Armenia (only Yerevan listed), but Cordoba is in the node set thanks to Argentina. Your answer should use a single XPath expression. Please store the results in a list `secondPops` of integers.

<!-- grade_id: cell-9cdf563633105d69, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c8e11dca9ab80a40, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q10:** With reference to the `topnames` dataset, please find all years where the top female name started with an M and the count was strictly larger than 50,000. Please navigate to the appropriate attribute, rather than returning a list of elements.

<!-- grade_id: cell-8ae34923d3f049b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-64822d9577e1ac2b, points: 2, grade: true, solution: false -->
```python
assert nodeset[0] == '1915'
assert len(nodeset) == 78
```